<a href="https://colab.research.google.com/github/hyalamanchi/hds5210-2023/blob/main/week06/week06_assignment_2.ipynb" target="_parent"><img src="https://colab.research.google.com/assets/colab-badge.svg" alt="Open In Colab"/></a>

# Week 6 Exercises

_McKinney 6.1_

There are multiple ways to solve the problems below.  You can use any one of several approaches.  For example, you can read CSV files using Pandas or the csv module.  Your score won't depend on which modules you choose to use unless explicitly noted below, but your programming style will still matter.

### 30.1 List of Allergies

In this GitHub repository, there is a file called `allergies.json` that contains a list of patient allergies.  You will need to download this [file from here](https://raw.githubusercontent.com/paulboal/hds5210-2023/main/week06/allergies.json) and then upload it into Google Colab to run these examples. It is taken from sample data provided by the EHR vendor, Epic, here: https://open.epic.com/Clinical/Allergy

Take some time to look at the structure of the file.  You can open it directly in Jupyter by clicking the _Home_ icon, then the _from_instructor_ folder, and then the _data_ folder.

Within the file, you'll see that it is a dictionary with many items in it.  One of those items is called `entry` and that item is a list of things.  You can tell that because the item name is immediately followed by an opening square bracket, signifying the start of a list.  It's line 11 of the file: `  "entry": [`

Write a function named `allergy_count(json_file)` that takes as one parameter the name of the JSON file and returns an integer number of entries in that file.  Your function should open the file, read the json into a Python object, and return how many items there are in the list of `entry`s.

```python
import json
ALLERGIES_FILE=json.load(open("allergies.json"))

```

<!-- grade_id: q1-answer, points: 0, grade: false, solution: true -->
```python
#
def allergy_count(data):
  """
  (dict) -> int
  Here first we assign patient list as 0 ,then we use for loop for patient indexing
  from the allergy file then we use if condition for the patient entries and the
  resource in the dictionary ,we checked here the status is equal to the confirmed
  then we are incrimenting patient_list to 1,after completion of for loop we
  are returning the patient_list then we are geting how many allergies are counted.

  >>> allergy_count(ALLERGIES_FILE) == 4
  True

  """
  patient_list=0
  for  patient_entries in range(len(ALLERGIES_FILE['entry'])):
    if ALLERGIES_FILE['entry'][patient_entries]['resource']['status']=='confirmed':
      patient_list=patient_list+1
  return patient_list
```

```python
allergy_count(ALLERGIES_FILE)
```

Output:
```
4
```

<!-- grade_id: q1-tests, points: 5, grade: true, solution: false -->
```python
assert type(allergy_count(ALLERGIES_FILE)) == int
assert allergy_count(ALLERGIES_FILE) == 4
```

```python
import doctest
doctest.run_docstring_examples(allergy_count,globals(),verbose=True)
```

Output:
```
PYDEV DEBUGGER WARNING:
sys.settrace() should not be used when the debugger is being used.
This may cause the debugger to stop working correctly.
If this is needed, please check: 
http://pydev.blogspot.com/2007/06/why-cant-pydev-debugger-work-with.html
to see how to restore the debug tracing back correctly.
Call Location:
  File "/usr/lib/python3.10/doctest.py", line 1501, in run
    sys.settrace(save_trace)

Finding tests in NoName
Trying:
    allergy_count(ALLERGIES_FILE) == 4
Expecting:
    True
ok
```

### 30.2 Number of Patients

If you dig a little bit deaper into this list of allergies, you'll see that each result has a patient associated with it.  Create a funcation called `patient_count(json_file)` that will count how many unique patients we have in this JSON structure.  

```python
# Put your solution here

def patient_count(data):
  """
  (dict)->int

  we create a empty patient_list and then by using for loop for entry indexing ALLERGIES_FILE  ,then append the
  patient_list ,we get the patient list then by using set we get the  unique  patient names,by returning the length of the
  number of patients.
  >>> patient_count(ALLERGIES_FILE)
  2

  """
  patient_list=[]
  for entries in range(len(ALLERGIES_FILE['entry'])):
    patient_list.append(ALLERGIES_FILE['entry'][entries]['resource']['patient']['display'])
  no_of_patients=set(patient_list)
  return len(no_of_patients)
```

```python
patient_count(ALLERGIES_FILE)
```

Output:
```
2
```

```python
assert type(patient_count(ALLERGIES_FILE)) == int
assert patient_count(ALLERGIES_FILE) == 2
```

```python
import doctest
doctest.run_docstring_examples(patient_count,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    patient_count(ALLERGIES_FILE)
Expecting:
    2
ok
```

### 30.3 How Many Allergies per Patient

Although each entry is a separate allergy, several of them are for the same patient.  Write a function called `allergy_per_patient(json_file)` that counts up how many allergies each patient has.


```python
# Put your solution here
def allergy_per_patient(data):
  """
  (dict)->dict
  Here we are creating the empty patient list and the allergy counts ,by using for loop entries in range length of
  allergies file ,then appending the allergies_file ,by doing set operation we are using two times for loop in first for
  loop we are using unique names ,and in inbuilt for loop we are using patient names in each entry so we inserting if
  condition that if unique name matches with the patient names in entry then count will be incrimented by 1.so we get
  the total counts of patients in the file. Now we are making a dictionary, where keys are patient names and values as
  total number of entries of that patient.
  >>> allergy_per_patient(ALLERGIES_FILE)
  {'Jason Argonaut': 3, 'Paul Boal': 1}
  """
  patient_list=[]
  allergy_counts={}
  for entries in range(len(ALLERGIES_FILE['entry'])):
    patient_list.append(ALLERGIES_FILE['entry'][entries]['resource']['patient']['display'])
    name_list=set(patient_list)
  for names in name_list:
    count=0
    for i in patient_list:
      if names==i:
        count+=1
    allergy_counts[names]=count
  return allergy_counts
```

```python
allergy_per_patient(ALLERGIES_FILE)
```

Output:
```
{'Jason Argonaut': 3, 'Paul Boal': 1}
```

```python
assert type(allergy_per_patient(ALLERGIES_FILE)) == dict
assert allergy_per_patient(ALLERGIES_FILE) == {'Paul Boal': 1, 'Jason Argonaut': 3}
```

```python
import doctest
doctest.run_docstring_examples(allergy_per_patient,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_per_patient(ALLERGIES_FILE)
Expecting:
    {'Jason Argonaut': 3, 'Paul Boal': 1}
ok
```

### 30.4 Patient Allergies and Reaction

You'll see in the file that each of the items in the `entry` list have several other attributes including a patient name, substance text representation, and a reaction manifestation.  Create a function named `allergy_list(json_file)` that will create an output list that has patient name, allergy, and reaction for each `entry`.  The actual result you should get will be:

```python
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis']]
```

You'll notice that the reaction and the manifestation of that action are lists.  You only need to capture the first reaction and the first manifestation of the action.  That is, if there is a list of things, just output the first one.

<!-- grade_id: q2-answer, points: 0, grade: false, solution: true -->
```python

# Put your solution here

def allergy_list(data):
  """
  (dict)-> list
  Here we are creating empty list, by assigning for loop to the entries in range from the allergies file,
  we are appending the list of patient, drug and reaction in each loop to the empty list.

  >>> allergy_list(ALLERGIES_FILE)
  [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]

 """

  patient_details=[]
  for entries in range (len(ALLERGIES_FILE['entry'])):
      patient=ALLERGIES_FILE['entry'][entries]['resource']['patient']['display']
      drug=ALLERGIES_FILE['entry'][entries]['resource']['substance']['text']
      reaction= ALLERGIES_FILE['entry'][entries]['resource']['reaction'][0]['manifestation'][0]['text']
      patient_details.append([patient,drug,reaction])
  return patient_details

```

```python
allergy_list(ALLERGIES_FILE)
```

Output:
```
[['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]
```

<!-- grade_id: q2-tests, points: 10, grade: true, solution: false -->
```python
assert allergy_list(ALLERGIES_FILE) == [['Jason Argonaut', 'PENICILLIN G', 'Hives'],
 ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'],
 ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'],
 ['Paul Boal', 'PENICILLIN G', 'Bruising']]

```

```python
import doctest
doctest.run_docstring_examples(allergy_list,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_list(ALLERGIES_FILE)
Expecting:
    [['Jason Argonaut', 'PENICILLIN G', 'Hives'], ['Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS', 'Itching'], ['Jason Argonaut', 'STRAWBERRY', 'Anaphylaxis'], ['Paul Boal', 'PENICILLIN G', 'Bruising']]
ok
```

### 30.5 Allergy Reaction

Write a function called `allergy_reaction(json_file,patient,substance)` that takes three parameter and returns the reaction that will happen if the patient takes the specified substance.  You can solve this, in part, by calling your `allergy_list` function inside your new `allergy_reaction` function.

If the substance is not found in the allergy list, the function should return None.

<!-- grade_id: q3-answer, points: 0, grade: false, solution: true -->
```python
# Put your solution here
def allergy_reaction(data,patient_name='name',drug='medicine'):
  """
  ('dict',str,str)->str
  Here intially we are making report as none ,then by using the for loop entries in range patient
  and drug as medicine ,later by using if condition we confirmed the patient_name and medicine he is allergic to.
  we call the allergic reaction as report. If patient and medicine details doesn't match with the data we had
  in our json file then the for loop ends and return report as None.
  >>> allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')
  'Bruising'

  """
  report=None
  for entries in range (len(ALLERGIES_FILE['entry'])):
      patient=ALLERGIES_FILE['entry'][entries]['resource']['patient']['display']
      medicine=ALLERGIES_FILE['entry'][entries]['resource']['substance']['text']
      if patient==patient_name and medicine==drug:
        report= ALLERGIES_FILE['entry'][entries]['resource']['reaction'][0]['manifestation'][0]['text']
        return report
  if report==None:
     return report
```

```python
allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G')
```

Output:
```
'Hives'
```

<!-- grade_id: q3-tests, points: 5, grade: true, solution: false -->
```python
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN G') == 'Hives'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'SHELLFISH-DERIVED PRODUCTS') == 'Itching'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'STRAWBERRY') == 'Anaphylaxis'
assert allergy_reaction(ALLERGIES_FILE, 'Jason Argonaut', 'PENICILLIN') == None
assert allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G') == 'Bruising'
```

```python
import doctest
doctest.run_docstring_examples(allergy_reaction,globals(),verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    allergy_reaction(ALLERGIES_FILE, 'Paul Boal', 'PENICILLIN G')
Expecting:
    'Bruising'
ok
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to save this notebook file back to GitHub.  To do that in Google Colab:
1. File -> Save a Copy in GitHub
2. Make sure your HDS5210 repository is selected
3. Make sure the file name includes the week number like this: `week06/week06_assignment_2.ipynb`
4. Add a commit message that means something

**Be sure week names are lowercase and use a two digit week number!!**

**Be sure you use the same file name provided by the instructor!!**

